[1507348887] SERVICE ALERT: HYDRA;Memory Usage;OK;HARD;3;Memory usage: total:37547.67 MB - used: 30036.78 MB (80%) - free: 7510.89 MB (20%)
[1507351287] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;1;Memory usage: total:37547.67 MB - used: 30094.31 MB (80%) - free: 7453.36 MB (20%)
[1507351407] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;2;Memory usage: total:37547.67 MB - used: 30053.46 MB (80%) - free: 7494.21 MB (20%)
[1507351527] SERVICE ALERT: HYDRA;Memory Usage;WARNING;HARD;3;Memory usage: total:37547.67 MB - used: 30067.82 MB (80%) - free: 7479.85 MB (20%)
[1507352727] SERVICE ALERT: HYDRA;Memory Usage;OK;HARD;3;Memory usage: total:37547.67 MB - used: 30036.90 MB (80%) - free: 7510.77 MB (20%)
[1507352727] SERVICE FLAPPING ALERT: HYDRA;Memory Usage;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1507353327] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;1;Memory usage: total:37547.67 MB - used: 30046.04 MB (80%) - free: 7501.63 MB (20%)
[1507353447] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;2;Memory usage: total:37547.67 MB - used: 30054.09 MB (80%) - free: 7493.58 MB (20%)
[1507353567] SERVICE ALERT: HYDRA;Memory Usage;WARNING;HARD;3;Memory usage: total:37547.67 MB - used: 30047.11 MB (80%) - free: 7500.56 MB (20%)
[1507364967] SERVICE FLAPPING ALERT: HYDRA;Memory Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1507374567] SERVICE ALERT: HYDRA;Memory Usage;OK;HARD;3;Memory usage: total:37547.67 MB - used: 30036.78 MB (80%) - free: 7510.89 MB (20%)
[1507375167] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;1;Memory usage: total:37547.67 MB - used: 30056.28 MB (80%) - free: 7491.39 MB (20%)
[1507375287] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;2;Memory usage: total:37547.67 MB - used: 30047.86 MB (80%) - free: 7499.81 MB (20%)
[1507375407] SERVICE ALERT: HYDRA;Memory Usage;WARNING;HARD;3;Memory usage: total:37547.67 MB - used: 30047.58 MB (80%) - free: 7500.09 MB (20%)
[1507403258] SERVICE ALERT: HYDRA-plex;Plex Media server;CRITICAL;SOFT;1;Plex Media Server.exe: not running
[1507403305] SERVICE ALERT: HYDRA-plex;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1507403335] HOST ALERT: HYDRA-plex;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1507403348] SERVICE ALERT: HYDRA-plex;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout
[1507403351] HOST ALERT: HYDRA-plex;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.1.20)
[1507403415] SERVICE ALERT: HYDRA-plex;CPU Load;OK;SOFT;2;CPU Load 13% (5 min average)
[1507403415] HOST ALERT: HYDRA-plex;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.89 ms
[1507403458] SERVICE ALERT: HYDRA-plex;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 1 minute(s)
[1507403607] SERVICE ALERT: HYDRA;Memory Usage;OK;HARD;3;Memory usage: total:37547.67 MB - used: 29820.14 MB (79%) - free: 7727.53 MB (21%)
[1507403978] SERVICE ALERT: HYDRA-plex;Plex Media server;OK;SOFT;3;Plex Media Server.exe: Running
[1507428207] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;1;Memory usage: total:37547.67 MB - used: 30043.18 MB (80%) - free: 7504.49 MB (20%)
[1507428327] SERVICE ALERT: HYDRA;Memory Usage;OK;SOFT;2;Memory usage: total:37547.67 MB - used: 30002.24 MB (80%) - free: 7545.43 MB (20%)