[1507459527] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;1;Memory usage: total:37547.67 MB - used: 30323.55 MB (81%) - free: 7224.12 MB (19%)
[1507459647] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;2;Memory usage: total:37547.67 MB - used: 30313.85 MB (81%) - free: 7233.82 MB (19%)
[1507459767] SERVICE ALERT: HYDRA;Memory Usage;WARNING;HARD;3;Memory usage: total:37547.67 MB - used: 30074.14 MB (80%) - free: 7473.53 MB (20%)
[1507462767] SERVICE ALERT: HYDRA;Memory Usage;OK;HARD;3;Memory usage: total:37547.67 MB - used: 29904.47 MB (80%) - free: 7643.20 MB (20%)
[1507478367] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;1;Memory usage: total:37547.67 MB - used: 30051.19 MB (80%) - free: 7496.48 MB (20%)
[1507478487] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;2;Memory usage: total:37547.67 MB - used: 30051.29 MB (80%) - free: 7496.38 MB (20%)
[1507478607] SERVICE ALERT: HYDRA;Memory Usage;WARNING;HARD;3;Memory usage: total:37547.67 MB - used: 30057.00 MB (80%) - free: 7490.67 MB (20%)
[1507479207] SERVICE ALERT: HYDRA;Memory Usage;OK;HARD;3;Memory usage: total:37547.67 MB - used: 29921.00 MB (80%) - free: 7626.67 MB (20%)
[1507486407] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;1;Memory usage: total:37547.67 MB - used: 30180.87 MB (80%) - free: 7366.80 MB (20%)
[1507486527] SERVICE ALERT: HYDRA;Memory Usage;OK;SOFT;2;Memory usage: total:37547.67 MB - used: 29887.24 MB (80%) - free: 7660.43 MB (20%)
[1507511727] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;1;Memory usage: total:37547.67 MB - used: 30085.69 MB (80%) - free: 7461.98 MB (20%)
[1507511847] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;2;Memory usage: total:37547.67 MB - used: 30061.16 MB (80%) - free: 7486.51 MB (20%)
[1507511967] SERVICE ALERT: HYDRA;Memory Usage;WARNING;HARD;3;Memory usage: total:37547.67 MB - used: 30051.75 MB (80%) - free: 7495.92 MB (20%)
[1507512567] SERVICE ALERT: HYDRA;Memory Usage;OK;HARD;3;Memory usage: total:37547.67 MB - used: 29894.10 MB (80%) - free: 7653.57 MB (20%)
[1507516767] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;1;Memory usage: total:37547.67 MB - used: 30050.90 MB (80%) - free: 7496.77 MB (20%)
[1507516887] SERVICE ALERT: HYDRA;Memory Usage;WARNING;SOFT;2;Memory usage: total:37547.67 MB - used: 30053.84 MB (80%) - free: 7493.83 MB (20%)
[1507517007] SERVICE ALERT: HYDRA;Memory Usage;WARNING;HARD;3;Memory usage: total:37547.67 MB - used: 30053.85 MB (80%) - free: 7493.82 MB (20%)